# <span style="color:blue">Programming for Data Science - DS-GA 1007</span>
## <span style="color:blue">Assignment 2</span>
### Yuqi Dai

## 1. Function return value

This exercise is to show you how function return values work. Run the code below to see what happens.

```python
def f1(x):
    print(x + 1)

def f2(x):
    return(x + 1)

f1(4)
f2(4)
```

Output:
```
5
5
```

Each function should produce the same output. Now run the following code.

```python
f1(4) + 1
f2(4) + 1
```

Output:
```
5
TypeError                                 Traceback (most recent call last)
<ipython-input-2-e43349f4d454> in <module>()
----> 1 f1(4) + 1
      2 f2(4) + 1

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

__Question__: Explain what is happening and why you get the error message.

<span style="color:gray">(Place your answer in this markdown cell)</span><br>
__Answer__:When the function is called, it will return nothing which is a NoneType object and 1 is an interger, so we cannot add interger to a nonetype object.

```python
type(f1(3))
```

Output:
```
4
NoneType
```

---
## 2. Function operation

Consider the following piece of code:

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
```

__Question__: Describe, in words, what this code does, and how it does it.

<span style="color:gray">(Place your answer in this markdown cell)</span><br>
__Answer:__ It returns the nth term of a Fibonacci sequence. It defines the nth term as the sum of two previous terms and set the first term to be 1 and 0th term to be 0.

__Question__: For what inputs will this function not behave as expected? What will happen?

<span style="color:gray">(Place your answer in this markdown cell)</span><br>
__Answer:__ If we set negative number or a fraction, the recursion will not be ended because 0 or 1 cannot be reached by subtracting 1 and 2 from the input. So the maximal depth of recursion will be reached.

---
## 3. Defining a function

The following program plays the game of "Rock, Paper, and Scissors". Familiarize yourself with it in order to understand how it works. Try running the program.

```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    if user_choice == computer_choice:
        print("Tie Game")
    elif [user_choice, computer_choice] in winners:
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Output:
```
You are playing a Rock, Paper, Scissors simulation.
Select R for rock, P for Paper, S for scissors, Q to quit
> s
Computer chooses S
Tie Game
Select R for rock, P for Paper, S for scissors, Q to quit
> r
Computer chooses R
Tie Game
Select R for rock, P for Paper, S for scissors, Q to quit
> p
Computer chooses S
Computer wins!
Select R for rock, P for Paper, S for scissors, Q to quit
> q
```

Recall how we define a function using `def`, and how we pass in parameters. Transform the code into a function called `rps` that takes two parameters: `user_choice` and `computer_choice`, with valid inputs `'R'`, `'P'`, and `'S'`. Your function should return `'D'` for a tie, `'W'` if `user_choice` wins, and `'L'` if `user_choice` loses.

<!-- grade_id: Q3B, points: 0, grade: false, solution: true -->
```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

### BEGIN SOLUTION
def rps(user_choice,computer_choice):
    if user_choice==computer_choice:
        return "D"
    if [user_choice,computer_choice] in winners:
        return "W"
    else:
        return "L"
### END SOLUTION

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    result = rps(user_choice, computer_choice)
    if result == 'D':
        print("Tie Game")
    elif result == 'W':
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Output:
```
You are playing a Rock, Paper, Scissors simulation.
Select R for rock, P for Paper, S for scissors, Q to quit
> r
Computer chooses R
Tie Game
Select R for rock, P for Paper, S for scissors, Q to quit
> q
```

If your code is correct, the asserts below will result in no messages.

<!-- grade_id: Q3C, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(rps('R', 'R'), 'D')
assert_equal(rps('R', 'P'), 'L')
assert_equal(rps('R', 'S'), 'W')
assert_equal(rps('P', 'R'), 'W')
assert_equal(rps('P', 'P'), 'D')
assert_equal(rps('P', 'S'), 'L')
assert_equal(rps('S', 'R'), 'L')
assert_equal(rps('S', 'P'), 'W')
assert_equal(rps('S', 'S'), 'D')
```

---
## 4. List of Tuples

Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple. 

For example: `[(1, 7), (1, 3), (3, 4, 5), (2, 2)]` yields `[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`

<!-- grade_id: Q4A, points: 0, grade: false, solution: true -->
```python
def sort_tuple_list(tuple_list):
    '''Function to sort a list of non-empty tuples in increasing 
       order by the last element in each tuple
    '''
    ### BEGIN SOLUTION
    A=[]
    for i in tuple_list:
        A.append(i[-1])
        A.sort()
    B=[]
    j=0
    while len(B)<len(A):
        for i in tuple_list:
            if j==len(A):
                break
            if i[-1]==A[j]:
                B.append(i)
                j+=1
    return B
    ### END SOLUTION
```

<!-- grade_id: Q4B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sort_tuple_list([(1, 3), (3, 2), (2, 1)]), [(2, 1), (3, 2), (1, 3)])
assert_equal(sort_tuple_list([(2, 3), (1, 1), (3, 2)]), [(1, 1), (3, 2), (2, 3)])
assert_equal(sort_tuple_list([(1, 7), (1, 3), (3, 4, 5), (2, 2)]), [(2, 2), (1, 3), (3, 4, 5), (1, 7)])
```

---
## 5. Merging Sorted Lists

Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. Ideally, the solution should work in "linear" time, making a single pass of both lists.

<!-- grade_id: Q5A, points: 0, grade: false, solution: true -->
```python
import string

def linear_merge(list1, list2):
    '''Given two lists sorted in increasing order, create and 
       return a merged list of all the elements in sorted order
    '''
    ### BEGIN SOLUTION
    order=list(string.ascii_lowercase)
    
    A=[]
    m=0
    n=0

    for i in order: 
        while list1[m][0]==i:
            A.append(list1[m])
            if m<len(list1)-1:
                m+=1
            else:
                break
        while list2[n][0]==i:
            A.append(list2[n])
            if n<len(list2)-1:
                n+=1
            else:
                break
    return A
        
            
    ### END SOLUTION
```

<!-- grade_id: Q5B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(linear_merge(['aa', 'xx'], ['bb', 'cc', 'zz']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(linear_merge(['aa', 'aa'], ['aa', 'bb', 'bb']), ['aa', 'aa', 'aa', 'bb', 'bb'])
```

# 6. Built-in Operator

Imagine that Python doesn’t have the `!=` (not equal) operator built in. Write a function that takes two parameters and gives the same result as the `!=` operator. Obviously, you cannot use `!=` within your function! Test if your code works by thinking of examples and making sure the output is the same for your new method as `!=` gives you.

<!-- grade_id: Q6A, points: 0, grade: false, solution: true -->
```python
def not_equal(val1, val2):
    ''' User supplied version of the != operator, 
        implemeted without using !=
    '''
    ### BEGIN SOLUTION
    if type(val1)==type(val2):
        if val1==val2:
            return False
        else:
            return True
    else:
        return True
    ### END SOLUTION
```

<!-- grade_id: Q6B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true, assert_false
assert_true(not_equal(2,3))
assert_false(not_equal(3,3))
assert_true(not_equal('3',3))
assert_true(not_equal((1,2),[1,2]))
assert_false(not_equal((3,4),(3,4)))
```

## 7. Quadratic Formula

A quadratic equation is a second order polynomial with three coefficients a, b, and c. It is given by:

$$ax^2 + bx + c = 0$$

The solutions to the equation are given by the formulas:

$$x_1 = \frac{-b + \sqrt{b^2 - 4ac}}{2a}$$

and

$$x_2 = \frac{-b - \sqrt{b^2 - 4ac}}{2a}$$



Write a function that computes the roots of a quadratic equation and returns them as a tuple $(x_1,x_2)$. Your function should take three parameters, a, b, and c. If only one soluction exists, return $(x_1,None)$. If not solution exists or all coefficients are equal zero return $(None,None)$. Note that if

$$b^2 - 4ac < 0$$ 

then the roots will be complex. Hint: remember that Python has a built-in complex data type.

<!-- grade_id: Q7A, points: 0, grade: false, solution: true -->
```python
import math
def roots(a, b, c):
    ''' Computes the roots of the quadratic equation with 
        coefficients a, b, c, and returns them as a tuple (x1, x2)
    '''
    ### BEGIN SOLUTION
    a,b,c=float(a),float(b),float(c)
    
    if (a==0 and b==0 and c!=0) or a==0 and b==0 and c==0:
        return (None,None)
    elif a==0:
        x1=-c/b
        return (x1,None)
    
    elif (b**2-4*a*c)>=0:
        x1=(-b+(b**2-4*a*c)**(1/2))/(2*a)
        x2=(-b-(b**2-4*a*c)**(1/2))/(2*a)
        if x1==x2:
            return (x1,None)
        else:
            return (x1,x2)
    elif b**2-4*a*c<0:
        complexroot=complex(0,(4*a*c-b**2)**(1/2))/2*a
        return (-b/2*a+complexroot,-b/2*a-complexroot)
        
    ### END SOLUTION
```

<!-- grade_id: Q7B, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
assert_equal(roots(0,0,4), (None, None))
assert_equal(roots(0,1,2), (-2, None))
assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))
```

## 8. The Game of Nims 
(extra credit)

In this game, two players sit in front of a pile of 100 stones. They take turns, each removing between 1 and 5 stones (assuming there are at least 5 stones left in the pile). The person who removes the last stone(s) wins.

In this problem, you’ll write a function to play this game. It may seem tricky, so break it down into parts. Like many programs, we have to use nested loops (one loop inside another). In the outermost loop, we want to keep playing until we are out of stones. Inside that, we want to keep alternating players. You have the option of either writing two blocks of code, or keeping a variable that tracks the current player.

We are going to use the `input` function to obtain the user's move. We might want to have an innermost loop that checks if the user’s input is valid. Is it a number? Is it a valid number (e.g. between 1 and 5)? Are there enough stones in the pile to take off this many? If any of these answers are no, we should tell the user and re-ask them the question.

If you choose to write two blocks of code, the basic outline of the program should be something like this:

    while [pile is not empty]:
       while [player 1’s answer is not valid]:
          [ask player 1]
       [execute player 1’s move]
       [same as above for player 2]
       
Be careful with the validity checks. Specifically, we want to keep asking player 1 for their choice as long as their answer is not valid, BUT we want to make sure we ask them at least ONCE. So, for example, we will want to keep a variable that tracks whether their answer is valid, and set it to `False` initially.

There is no need to limit yourself to only one function! Use as many functions as you want in order to make the code more readable and modular.

When you're finished, test each other's programs by playing them!

<!-- grade_id: Q8A, points: 4, grade: true, solution: true -->
```python

def input1(n,x):
    while True:
        choice1=input('Player1, what is your choice?')
        rightinput=[]
        for i in range(1,x+1):
            rightinput.append(str(i))    
        if choice1 in rightinput:
            if int(choice1)<=n:
                return choice1
            else:
                print('Invalid choice, do it again please')
        else:
            print('Invalid choice, do it again please')

            
def input2(n,x):
    while True:
        choice2=input('Player2, what is your choice?')
        rightinput=[]
        for i in range(1,x+1):
            rightinput.append(str(i))
        if choice2 in rightinput:
            if int(choice2) <=n:
                return choice2
            else:
                print('Invalid choice, do it again please')
        else:
            print('Invalid choice, do it again please')
                
        
def play_nims(total_stone_number,max_number_stone_per_round):
    '''
    An interactive two-person game; also known as Stones.
    pile: the number of stones in the pile to start
    max_stones: the maximum number of stones you can 
    take in one turn
    '''
    print('Welcome to the stone pile challange, you can remove 1 to 5 stones each time and the player who removes the last two wins.')
   
    m=total_stone_number
    n=max_number_stone_per_round
    while True:
        if m<=5 and m>=2:
            print('There are {} stones left'.format(m))
            print('Congradulation! Player1 wins')
            break
        if m==1:
            print('There are {} stones left'.format(m))
            print('Tie!')
            break
        else: 
            print('There are {} stones left'.format(m))
        
        choice1=input1(m,n)
        
        m=m-int(choice1)
        
        
        
        if m<=5 and m>=2:
            print('There are {} stones left'.format(m))
            print('Congradulation! Player2 wins')
            break
        if m==1:
            print('There are {} stones left'.format(m))
            print('Tie!')
            break
        else: 
            print('There are {} stones left'.format(m))
        
        choice2=input2(m,n)
        
        m=m-int(choice2)
        
        

### GRADER INSTRUCTIONS
## Verify that their program correctly plays the game by running it. Try various values to ensure they are checking
## for valid inputs. 
## Award 1 mark for basic structure
## Award 1 mark if they validate input correctly
## Award 1 mark if the program correctly responds to invalid input by reasking
## Award 1 mark if they use more than one function
### END SOLUTION
```

```python
play_nims(100,5)
```

Output:
```
Welcome to the stone pile challange, you can remove 1 to 5 stones each time and the player who removes the last two wins.
There are 100 stones left
Player1, what is your choice?y
Invalid choice, do it again please
Player1, what is your choice?r
Invalid choice, do it again please
Player1, what is your choice?6
Invalid choice, do it again please
Player1, what is your choice?5
There are 95 stones left
Player2, what is your choice?5
There are 90 stones left
Player1, what is your choice?5
There are 85 stones left
Player2, what is your choice?5
There are 80 stones left
Player1, what is your choice?5
There are 75 stones left
Player2, what is your choice?5
There are 70 stones left
Player1, what is your choice?5
There are 65 stones left
Player2, what is your choice?
Invalid choice, do it again please
Player2, what is your choice?5
There are 60 stones left
Player1, what is your choice?5
There are 55 stones left
Player2, what is your choice?5
There are 50 stones left
Player1, what is your choice?
```